![DSA log](dsalogo.png)

### Instructions

1. Make sure you are using a version of notebook greater than v.3. If you installed Anaconda with python 3 - this is likely to be fine. The next piece of code will check if you have the right version.
2. The notebook has both some open test cases that you can use to test the functionality of your code - however it will be run on another set of test cases that you can't from which marks will be awarded. So passing all the tests in this notebook is not a guarantee that you have done things correctly - though its highly probable.
3. Also make sure you submit a notebook that doesn't return any errors. One way to ensure this is to run all the cells before you submit the notebook.
4. When you are done create a zip file of your notebook and upload that
5. For each cell where you see "YOUR CODE HERE" delete the return notImplemented statement when you write your code there - don't leave it in the notebook.
6. Once you are done, you are done.

```python

```

Output:
```
Collecting nose
  Downloading nose-1.3.7-py2-none-any.whl (154kB)
[K    100% |████████████████████████████████| 163kB 217kB/s ta 0:00:01
[?25hInstalling collected packages: nose
Successfully installed nose-1.3.7
[33mYou are using pip version 9.0.1, however version 9.0.3 is available.
You should consider upgrading via the 'pip install --upgrade pip' command.[0m
```

## Python Basics
<i>Ernest M (03/2018)</i>

```python
# Import some stuff that you will use in this notebook
from nose.tools import assert_equal
import numpy as np
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

#### Question 1

Complete the function to output a string which is the concatenation of the inputs to the function

<!-- grade_id: cell-9f15cc32a701e305, points: 0, grade: false, solution: true -->
```python
def get_string(x,y):
    return str(x) + str(y)
    raise NotImplementedError()
```

<!-- grade_id: cell-bd1e1e426e1025f2, points: 2, grade: true, solution: false -->
```python
assert_equal(get_string(4,5), '45')
assert_equal(get_string('hello','world'), 'helloworld')
```

#### Question 2

Complete the function to output the value of $y$ given $x$ and $p$. Function definition is $y = 4 + 5x^p$ <br>
**Condition:** If the $x$ is not positive the value of y should be 0

<!-- grade_id: cell-aefc8e898976745e, points: 0, grade: false, solution: true -->
```python
def ypolynomial(x,p):
    if x > 0:
        x = x**p
        y = 4 + 5*x
    else: 
        y = 0
    return y
    raise NotImplementedError()
```

<!-- grade_id: cell-c8503fd6f87f54f1, points: 3, grade: true, solution: false -->
```python
assert_equal(ypolynomial(3,3), 139)
assert_equal(ypolynomial(-2,3), 0)
```

#### Question 3

Complete the function to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $p$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

<!-- grade_id: cell-d096967be6b68c1f, points: 0, grade: false, solution: true -->
```python
def divisibles(N, p):
    result = []
    if N < 0:
        N=20
    else:
        pass
    
    i = 0
    while i <= N:
        if i%p == 0:
            result.append(i)
        else:
            pass
        i+=1
    return result
    raise NotImplementedError()
```

<!-- grade_id: cell-6cbc8b89f1bfbece, points: 3, grade: true, solution: false -->
```python
assert_equal(divisibles(10,2), [0,2,4,6,8,10])
assert_equal(divisibles(-10,2), [0,2,4,6,8,10,12,14,16,18,20])
```

<!-- grade_id: cell-c3a48b75b885a567, points: 3, grade: true, solution: false -->
```python
assert_equal(divisibles(6,4), [0,4])
```

#### Question 4
Re-do Question 3, but in this case instead of evaluating numbers starting from 0, let them start from $s$. Complete the function below where $s$ is the start of the numbers ending at $N$ inclusive.

<!-- grade_id: cell-58166f2ffb8d4331, points: 0, grade: false, solution: true -->
```python
def divisibles2(s, N, p):
    result = []
    if N < 0:
        N=20
    else:
        pass
    
    i = s
    while i <= N:
        if i%p == 0:
            result.append(i)
        else:
            pass
        i+=1
    return result
    raise NotImplementedError()
```

<!-- grade_id: cell-a517c73c87571847, points: 3, grade: true, solution: false -->
```python
assert_equal(divisibles2(0,10,2), [0,2,4,6,8,10])
assert_equal(divisibles2(2,-10,2), [2,4,6,8,10,12,14,16,18,20])
```

<!-- grade_id: cell-025d1639d2eff09f, points: 4, grade: true, solution: false -->
```python
assert_equal(divisibles2(4,5,2), [4])
```

#### Question 5

Complete the code to return a list of the first $N$ numbers that are divisible by 5,7 and 11

<!-- grade_id: cell-33dbd87e8be70627, points: 0, grade: false, solution: true -->
```python
def divisibo(N):
    result = []
    divs = [5, 7, 11]
    start = divs[2]
    
    while(len(result) < N):
        if start%11 == 0:
            if start%7 == 0:
                if start%5 == 0:
                    result.append(start)
        start+=1
    return result
    raise NotImplementedError()
```

<!-- grade_id: cell-1c7dd28aec0d15c8, points: 4, grade: true, solution: false -->
```python
assert_equal(divisibo(1), [385])
assert_equal(divisibo(2), [385, 770])
```

<!-- grade_id: cell-e7b3abc2fa705436, points: 5, grade: true, solution: false -->
```python
assert_equal(divisibo(4), [385, 770, 1155, 1540])
```

```python
divisibo(8)
```

Output:
```
[385, 770, 1155, 1540, 1925, 2310, 2695, 3080]
```

#### Question 6

Complete the function to output all primes up to the number $n$

<!-- grade_id: cell-d1ceaad5daa87eaa, points: 0, grade: false, solution: true -->
```python
def primes_upto(N):
    result = []
    if N > 1:
        result.append(2)
        i = 2
        while(i <= N):
            if(i%2 != 0):
                result.append(i)
            i+=1
        return result
    else:
        return result
    raise NotImplementedError()
```

```python
primes_upto(3)
```

Output:
```
[2, 3]
```

<!-- grade_id: cell-233f36d9259b2c92, points: 4, grade: true, solution: false -->
```python
assert_equal(primes_upto(2), [2])
assert_equal(primes_upto(3), [2,3])
```

<!-- grade_id: cell-b164a145610833fe, points: 4, grade: true, solution: false -->
```python
assert_equal(primes_upto(1), [])
```

#### Question 7

Complete the function to output the first $n$ primes

<!-- grade_id: cell-8e763c5dfbdc7ace, points: 0, grade: false, solution: true -->
```python
def primes_first(N):
    result = [2]
    i = 2
    while(len(result) < N):
        if(i%2 != 0):
            result.append(i)
        i+=1
    return result
    raise NotImplementedError()
```

<!-- grade_id: cell-28b5645753ea24e2, points: 3, grade: true, solution: false -->
```python
assert_equal(primes_first(1), [2])
assert_equal(primes_first(2), [2,3])
```

<!-- grade_id: cell-e14b769f50493306, points: 4, grade: true, solution: false -->
```python
assert_equal(primes_first(4), [2, 3, 5, 7])
```

#### Question 8

Complete the function to output a sorted list of elements of sublists of a list (confusing right ?) <br>
Example: given [[2,13,44], [6,7]] return [2,6,7,13,44]

<!-- grade_id: cell-a3301b8604a320f0, points: 0, grade: false, solution: true -->
```python
def flatten_list(superlist):
    # Flatten array
    sort = []
    i = 0
    while i < len(superlist):
        z = 0
        while(z < len(superlist[i])):
            sort.append(superlist[i][z])
            z+=1
        i+=1
    # Bubble sort
    for p in range(len(sort)-1,0,-1):
        for x in range(p):
            if sort[x]>sort[x+1]:
                tmp = sort[x]
                sort[x] = sort[x+1]
                sort[x+1] = tmp
    return sort
    raise NotImplementedError()
```

<!-- grade_id: cell-9f7399759b36c197, points: 5, grade: true, solution: false -->
```python
assert_equal(flatten_list([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
assert_equal(flatten_list([[2], [61,34,5,8,9]]), [2, 5, 8, 9, 34, 61])
```

#### Question 9

Complete the function to return the longest string in a line of text <br>
Example: for text "Hallo people of the world!" should return "people" <br>
If there is a tie, then it should return the 1st "longest" word encountered

<!-- grade_id: cell-8bb7222f98203461, points: 0, grade: false, solution: true -->
```python
def longest_word(text):
    words = text.split()
    largest = 0
    for word in words:
        current_len = len(word)
        if (current_len > largest):
            final_word = word
            largest = current_len
        else:
            largest = largest
    return final_word
    raise NotImplementedError()
```

<!-- grade_id: cell-ab4e86977747d1d2, points: 2, grade: true, solution: false -->
```python
assert_equal(longest_word("Hello there people"), 'people')
assert_equal(longest_word("such a great day today "), 'great')
```

<!-- grade_id: cell-130ea877bc75c47a, points: 3, grade: true, solution: false -->
```python
assert_equal(longest_word("toa maneno yako"), 'maneno')
```

#### Question 10

Write a function that takes a list as input and returns a dictionary with the elements of the list as the keys and the frequency of the element as the values

<!-- grade_id: cell-762d60e0a18c5d08, points: 0, grade: false, solution: true -->
```python
def dictionarize(alist):
    diction = {}
    for x in alist:
        i = alist.count(x)
        if i > 0:
            diction[x] = i
    return diction
    raise NotImplementedError()
```

<!-- grade_id: cell-bfa722edbbb1ef66, points: 3, grade: true, solution: false -->
```python
assert_equal(dictionarize([1,2,2,3,4,5]),{1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
assert_equal(dictionarize(list('abccdeedeecca')), {'a': 2, 'b': 1, 'c': 4, 'd': 2, 'e': 4})
```

<!-- grade_id: cell-a058708897a7ded4, points: 5, grade: true, solution: false -->
```python
assert_equal(dictionarize([1,'a','h',5]),{1: 1, 'h': 1, 5: 1, 'a': 1})
```

```python

```